// variables3.rs
// Make me compile! Execute the command `rustlings hint variables3` if you want a hint :)

// I AM NOT DONE

fn main() {
    let mut x = 3; // solution: needs to be declared as mutable
    println!("Number {}", x);
    x = 5; // don't change this line
    println!("Number {}", x);
}
